# Module 6 Assignment

A few things you should keep in mind when working on assignments:

1. Run the first code cell to import modules needed by this assignment before proceeding to problems.
2. Make sure you fill in any place that says `# YOUR CODE HERE`. Do not write your answer anywhere else other than where it says `# YOUR CODE HERE`. Anything you write elsewhere will be removed or overwritten by the autograder.
3. Each problem has an autograder cell below the answer cell. Run the autograder cell to check your answer. If there's anything wrong in your answer, the autograder cell will display error messages.
4. Before you submit your assignment, make sure everything runs as expected. Go to the menubar, select Kernel, and Restart & Run all. If the notebook runs through the last code cell without an error message, you've answered all problems correctly.
5. Make sure that you save your work (in the menubar, select File → Save and CheckPoint).

-----

# Run Me First!

```python
import pandas as pd
import numpy as np
import seaborn as sns

from nose.tools import assert_equal, assert_almost_equal, assert_true, assert_is_instance

# We do this to ignore warnings
import warnings
warnings.filterwarnings("ignore")
```

---
# Prepare Yelp Review Data

In this assignment, you will use the Yelp review dataset. Before we attempt to build a model, we first prepare the data.

The Yelp review dataset contains 1000 customer reviews of a group of restaurants. There are two columns in the dataset, column **stars** which is the star rating and column __text__ which is the review text. The dataset only contains 1-star and 5-star reviews. There are 500 1-star reviews and 500 5-star reviews.

Please run the next Code cell before proceeding to Problem 1.


```python
#Load yelp review dataset
df = pd.read_csv('data/yelp_reviews.csv')
rating = df['stars']
data = df['text']
sample_review = data[0]
print(f'Counts of Reviews by Star Rating:\n{rating.value_counts()}\n')
print(f'Sample Review({rating[0]} stars):\n{sample_review}')
print('-'*80)
print('\nDataset Sample:')
df.sample(2, random_state=2)
```

Output:
```
Counts of Reviews by Star Rating:
5    500
1    500
Name: stars, dtype: int64

Sample Review(5 stars):
I love love LOVE this place. My boss (who is into healthy eating) recommended this place. I went over with some highly skeptical friends and one dinner was enough to convert them into believers! The food here is so good! We had the Shrimp dumplings and the Onion tart as starters. We ordered the Shirataki noodles and street tacos as entrees. So also ordered the Kale-aid. All of the dishes were yummy. 
I have gone back many times since then and have never been disappointed! I have gone after yoga to get some Kale salad or the chicken chopped salad. I always have to get the Kale aid. 
Once, a guy at the next table, uprooted a whole plant by mistake (on the patio) and was highly embarrassed as was his date! Ever since, I have very careful not to throw my arms around as I can be quite clumsy sometimes! I do NOT want to be banned from my favorite place for my clumsiness! I don't think I can
```

---
# Problem 1: Get Top Words from a Sample Review

Get the most-used words in a sample review.

For this problem you will use the string **sample_review** created above.

To solve this problem do the following:
- Use regular expression `r'[^\w]'` to replace all characters in sample_review that are not letter or number with white spaces, assign the new string to sample_review_ns.
- Convert sample_review_ns to **all lower case** with string function lower(), assign the new string to sample_review_nsl.
- Use string function split() to split sample_review_nsl to a list of words and use the list to construct a `Counter` object in `collections` module.
- Get **top 10** words from the Counter object with `most_common()` function and assign the result to variable __top_10_words__.

After this problem, there's a new variable **top_10_words** defined.

---

<!-- grade_id: p1-answer, points: 0, grade: false, solution: true -->
```python
import re
import collections as cl

# YOUR CODE HERE
sample_review_ns = re.sub(r'[^\w]',' ',sample_review)
```

```python
sample_review_ns
```

Output:
```
'I love love LOVE this place  My boss  who is into healthy eating  recommended this place  I went over with some highly skeptical friends and one dinner was enough to convert them into believers  The food here is so good  We had the Shrimp dumplings and the Onion tart as starters  We ordered the Shirataki noodles and street tacos as entrees  So also ordered the Kale aid  All of the dishes were yummy   I have gone back many times since then and have never been disappointed  I have gone after yoga to get some Kale salad or the chicken chopped salad  I always have to get the Kale aid   Once  a guy at the next table  uprooted a whole plant by mistake  on the patio  and was highly embarrassed as was his date  Ever since  I have very careful not to throw my arms around as I can be quite clumsy sometimes  I do NOT want to be banned from my favorite place for my clumsiness  I don t think I can live without True Food '
```

```python
sample_review_nsl = sample_review_ns.lower()
```

```python
sample_review_nsl
```

Output:
```
'i love love love this place  my boss  who is into healthy eating  recommended this place  i went over with some highly skeptical friends and one dinner was enough to convert them into believers  the food here is so good  we had the shrimp dumplings and the onion tart as starters  we ordered the shirataki noodles and street tacos as entrees  so also ordered the kale aid  all of the dishes were yummy   i have gone back many times since then and have never been disappointed  i have gone after yoga to get some kale salad or the chicken chopped salad  i always have to get the kale aid   once  a guy at the next table  uprooted a whole plant by mistake  on the patio  and was highly embarrassed as was his date  ever since  i have very careful not to throw my arms around as i can be quite clumsy sometimes  i do not want to be banned from my favorite place for my clumsiness  i don t think i can live without true food '
```

```python
words = sample_review_nsl.split()
```

```python
print(words)
```

Output:
```
['i', 'love', 'love', 'love', 'this', 'place', 'my', 'boss', 'who', 'is', 'into', 'healthy', 'eating', 'recommended', 'this', 'place', 'i', 'went', 'over', 'with', 'some', 'highly', 'skeptical', 'friends', 'and', 'one', 'dinner', 'was', 'enough', 'to', 'convert', 'them', 'into', 'believers', 'the', 'food', 'here', 'is', 'so', 'good', 'we', 'had', 'the', 'shrimp', 'dumplings', 'and', 'the', 'onion', 'tart', 'as', 'starters', 'we', 'ordered', 'the', 'shirataki', 'noodles', 'and', 'street', 'tacos', 'as', 'entrees', 'so', 'also', 'ordered', 'the', 'kale', 'aid', 'all', 'of', 'the', 'dishes', 'were', 'yummy', 'i', 'have', 'gone', 'back', 'many', 'times', 'since', 'then', 'and', 'have', 'never', 'been', 'disappointed', 'i', 'have', 'gone', 'after', 'yoga', 'to', 'get', 'some', 'kale', 'salad', 'or', 'the', 'chicken', 'chopped', 'salad', 'i', 'always', 'have', 'to', 'get', 'the', 'kale', 'aid', 'once', 'a', 'guy', 'at', 'the', 'next', 'table', 'uprooted', 'a', 'whole', 'plant', 'by', 'mistak
```

```python
wc = cl.Counter(words)
```

```python
print("Word Count:\n", wc)
```

Output:
```
Word Count:
 Counter({'i': 10, 'the': 10, 'and': 5, 'to': 5, 'have': 5, 'my': 4, 'as': 4, 'love': 3, 'place': 3, 'was': 3, 'kale': 3, 'this': 2, 'is': 2, 'into': 2, 'some': 2, 'highly': 2, 'food': 2, 'so': 2, 'we': 2, 'ordered': 2, 'aid': 2, 'gone': 2, 'since': 2, 'get': 2, 'salad': 2, 'a': 2, 'not': 2, 'can': 2, 'be': 2, 'boss': 1, 'who': 1, 'healthy': 1, 'eating': 1, 'recommended': 1, 'went': 1, 'over': 1, 'with': 1, 'skeptical': 1, 'friends': 1, 'one': 1, 'dinner': 1, 'enough': 1, 'convert': 1, 'them': 1, 'believers': 1, 'here': 1, 'good': 1, 'had': 1, 'shrimp': 1, 'dumplings': 1, 'onion': 1, 'tart': 1, 'starters': 1, 'shirataki': 1, 'noodles': 1, 'street': 1, 'tacos': 1, 'entrees': 1, 'also': 1, 'all': 1, 'of': 1, 'dishes': 1, 'were': 1, 'yummy': 1, 'back': 1, 'many': 1, 'times': 1, 'then': 1, 'never': 1, 'been': 1, 'disappointed': 1, 'after': 1, 'yoga': 1, 'or': 1, 'chicken': 1, 'chopped': 1, 'always': 1, 'once': 1, 'guy': 1, 'at': 1, 'next': 1, 'table': 1, 'uprooted': 1, 'whole':
```

```python
print("\n Top 10 Words:")
top_10_words = wc.most_common(10)
print(top_10_words)
```

Output:
```
Top 10 Words:
[('i', 10), ('the', 10), ('and', 5), ('to', 5), ('have', 5), ('my', 4), ('as', 4), ('love', 3), ('place', 3), ('was', 3)]
```

<!-- grade_id: p1-test, points: 1, grade: true, solution: false -->
```python
assert_equal(type(top_10_words), list, msg='top_10_words should be a list')
assert_equal(top_10_words, [('i', 10),('the', 10),('and', 5),('to', 5),('have', 5),('my', 4),
                            ('as', 4),('love', 3),('place', 3),('was', 3)], msg='top_10_words list is not correct')
print('Top 10 Words:')
top_10_words
```

Output:
```
Top 10 Words:
[('i', 10),
 ('the', 10),
 ('and', 5),
 ('to', 5),
 ('have', 5),
 ('my', 4),
 ('as', 4),
 ('love', 3),
 ('place', 3),
 ('was', 3)]
```

---
# Problem 2: Prepare Data for Text Classification

Create label and split text data set to training and testing set.

For this problem you will use **data** and __rating__ created above.

To solve this problem do the following:
- Create **label** from __rating__. Set label to 1 if rating is 5 stars and 0 otherwise.(Hint: use lambda function: `label = rating.map(lambda x: 1 if x==5 else 0)`
- Split data and label to training and testing using `train_test_split`.
 - Set test_size to 0.4.
 - Set random_state to 23.
 - Assign return values to **d_train, d_test, l_train, l_test**.

After this problem, there are four new variable, **d_train, d_test, l_train** and  __l_test__ defined.

---

```python
rating
```

Output:
```
0      5
1      5
2      5
3      5
4      5
      ..
995    1
996    1
997    1
998    1
999    1
Name: stars, Length: 1000, dtype: int64
```

```python
label = rating.map(lambda x: 1 if x==5 else 0)
```

```python
label
```

Output:
```
0      1
1      1
2      1
3      1
4      1
      ..
995    0
996    0
997    0
998    0
999    0
Name: stars, Length: 1000, dtype: int64
```

<!-- grade_id: p2-answer, points: 0, grade: false, solution: true -->
```python
from sklearn.model_selection import train_test_split

# YOUR CODE HERE
d_train, d_test, l_train, l_test = train_test_split(data, label, test_size=0.4, random_state=23)
```

<!-- grade_id: p2-test, points: 1, grade: true, solution: false -->
```python
assert_equal(len(d_train), 600, msg='training set size is not correct')
assert_equal(len(d_test), 400, msg='testing set size is not correct')
assert_equal(l_train.iloc[1], 1, msg='l_train is not correct')
assert_equal(l_train.iloc[100], 1, msg='l_train is not correct')
assert_equal(l_test.iloc[0], 0, msg='l_test is not correct')
assert_equal(l_test.iloc[100], 1, msg='l_test is not correct')
```

---

# Problem 3: Train a Multinomial Naive Bayes Model

For this problem, use d_train, l_train, d_test and l_test created above.

Your task for this problem is to build and train a `MultinomialNB` estimator to make predictions on the Yelp review dataset. 

To solve this problem do the following:
- Create a `TfidfVectorizer` object, set `stop_words` to 'english'.
- Fit the `TfidfVectorizer` objec with d_train
- Transform d_train and d_test with the `TfidfVectorizer` object to get **train_dtm** and __test_dtm__
- Create a `MultinomialNB` estimator **nb_model**. Accept default values for all hyperparameters.
- Fit the `MultinomialNB` estimator using train_dtm and l_train.

After this problem, there will be a trained multinomial Naive Bayes model **nb_model** define, as well as two document term matrices, __train_dtm, test_dtm__.

-----

<!-- grade_id: p3-answer, points: 0, grade: false, solution: true -->
```python
from sklearn.feature_extraction.text import TfidfVectorizer
from sklearn.naive_bayes import MultinomialNB

# YOUR CODE HERE
tfidf = TfidfVectorizer(stop_words="english")
```

```python
tfidf.fit(d_train)
```

Output:
```
TfidfVectorizer(analyzer='word', binary=False, decode_error='strict',
                dtype=<class 'numpy.float64'>, encoding='utf-8',
                input='content', lowercase=True, max_df=1.0, max_features=None,
                min_df=1, ngram_range=(1, 1), norm='l2', preprocessor=None,
                smooth_idf=True, stop_words='english', strip_accents=None,
                sublinear_tf=False, token_pattern='(?u)\\b\\w\\w+\\b',
                tokenizer=None, use_idf=True, vocabulary=None)
```

```python
train_dtm = tfidf.transform(d_train)
```

```python
train_dtm
```

Output:
```
<600x7361 sparse matrix of type '<class 'numpy.float64'>'
	with 30310 stored elements in Compressed Sparse Row format>
```

```python
test_dtm = tfidf.transform(d_test)
```

```python
test_dtm
```

Output:
```
<400x7361 sparse matrix of type '<class 'numpy.float64'>'
	with 16990 stored elements in Compressed Sparse Row format>
```

```python
nb_model = MultinomialNB()
```

```python
nb_model.fit(train_dtm,l_train)
```

Output:
```
MultinomialNB(alpha=1.0, class_prior=None, fit_prior=True)
```

<!-- grade_id: p3-test, points: 1, grade: true, solution: false -->
```python
assert_equal(type(nb_model), type(MultinomialNB()), msg="nb_model is not defined as a MultinomialNB model")
assert_almost_equal(train_dtm[0, 132], 0.05808376842859404, msg="train_dtm is not correct")
assert_almost_equal(test_dtm[0, 157], 0.12459374859215833, msg="test_dtm is not correct")

```

---

# Problem 4: Get Text Classification Metrics

For this problem, you will compute the classification metrics of the nb_model created in problem 3.  

To complete this function, you must explicitly:

- Apply nb_model `predict` function on test_dtm created in problem 3 to get predicted label, assign it to variable **l_pred**.
- Use l_test and l_pred to calculate:
 - The mean accuracy score using `accuracy_score` function in `metrics` module and save the value to **mas_score**.
 - The classification report using `classification_report` function in `metrics` module and save it to variable **c_report**.

After this problem, there will be two new variables, **mas_score** and __c_report__ defined.

-----

<!-- grade_id: p4-answer, points: 0, grade: false, solution: true -->
```python
from sklearn import metrics

# YOUR CODE HERE
l_pred = nb_model.predict(test_dtm)
```

```python
l_pred
```

Output:
```
array([0, 0, 1, 0, 0, 0, 0, 1, 0, 1, 1, 1, 0, 1, 1, 1, 0, 0, 0, 1, 0, 1,
       0, 1, 0, 0, 0, 0, 0, 1, 0, 1, 1, 1, 1, 1, 0, 0, 0, 1, 1, 1, 0, 0,
       0, 0, 0, 0, 0, 0, 0, 1, 1, 1, 1, 0, 0, 1, 0, 1, 0, 0, 1, 0, 1, 0,
       0, 0, 1, 0, 0, 1, 1, 1, 1, 1, 0, 1, 1, 1, 0, 1, 0, 1, 0, 0, 0, 0,
       0, 0, 1, 0, 1, 1, 0, 0, 0, 1, 1, 0, 1, 1, 0, 0, 0, 1, 0, 0, 1, 1,
       0, 1, 0, 1, 1, 0, 1, 1, 1, 0, 0, 0, 0, 1, 1, 1, 1, 0, 1, 1, 0, 0,
       1, 0, 0, 0, 0, 1, 0, 1, 1, 0, 0, 1, 0, 0, 0, 0, 0, 1, 0, 0, 0, 0,
       0, 0, 1, 1, 1, 0, 1, 0, 0, 0, 1, 1, 0, 1, 0, 1, 0, 0, 0, 1, 1, 1,
       1, 0, 1, 0, 1, 1, 0, 0, 0, 0, 0, 1, 0, 0, 1, 0, 0, 1, 0, 0, 1, 1,
       1, 0, 1, 1, 1, 0, 1, 0, 0, 0, 1, 0, 0, 0, 1, 1, 1, 1, 0, 1, 0, 1,
       1, 0, 0, 1, 0, 1, 1, 1, 1, 0, 0, 0, 0, 1, 0, 0, 0, 1, 0, 0, 1, 0,
       1, 0, 1, 0, 0, 1, 1, 0, 0, 1, 1, 1, 1, 1, 0, 1, 0, 0, 0, 1, 0, 0,
       1, 0, 0, 1, 0, 0, 1, 0, 1, 0, 0, 0, 1, 0, 0, 0, 0, 1, 1, 1, 1, 0,
       1, 1, 1, 1, 0, 0, 1, 0, 1, 1, 1, 0, 1, 0, 1,
```

```python
mas_score = metrics.accuracy_score(l_test,l_pred)
mas_score
```

Output:
```
0.885
```

```python
c_report = metrics.classification_report(l_test,l_pred)
```

<!-- grade_id: p4-test, points: 1, grade: true, solution: false -->
```python
assert_equal(mas_score, 0.885, msg='Accuracy score is not correct')
assert_true('0.94' in c_report, msg='Classification report is not correct')
assert_true('194' in c_report, msg='Classification report is not correct')
print(f"Accuracy Score: {mas_score*100:4.1f}%")
print(f"Classification Report\n{c_report}")
```

Output:
```
Accuracy Score: 88.5%
Classification Report
              precision    recall  f1-score   support

           0       0.84      0.94      0.89       194
           1       0.94      0.83      0.88       206

    accuracy                           0.89       400
   macro avg       0.89      0.89      0.88       400
weighted avg       0.89      0.89      0.88       400
```

```python

```